---
# Prometheus official unit-testing documentation - https://prometheus.io/docs/prometheus/latest/configuration/unit_testing_rules/
# rule_files contains the list of files to be tested
rule_files:
  - /tmp/rules.verify

# group_eval_order contains the list of groups to be tested
group_eval_order:
  - memcached.rules

tests:
# for each time frame based on the interval, we define the metrics values
  - interval: 1m
    input_series:
      - series: 'memcached_deployment_size_undesired_count_total'
        # time:  0 1 2 3 4 5 6 7 8 9 10 11 12 13 14 15
        values: "0 0 0 1 2 3 3 3 3 3  3  3  3  4  5  6"
      - series: 'memcached_operator_up_total'
        # time:  0 1 2 3 4 5 6 7 8 9 10 11 12 13 14 15
        values: "0 0 0 0 0 0 1 1 1 1  0  0  0  0  0  0"

# then, we evaluate the alerts behaviour in the eval_time we choose
    alert_rule_test:
      # it must not trigger before 5m
      - eval_time: 4m
        alertname: MemcachedDeploymentSizeUndesired
        exp_alerts: []
      - eval_time: 4m
        alertname: MemcachedOperatorDown
        exp_alerts: []
      # it must trigger after 5m
      - eval_time: 5m
        alertname: MemcachedDeploymentSizeUndesired
        exp_alerts:
          - exp_annotations:
              description: "Memcached-sample deployment size was not as desired more than 3 times in the last 5 minutes."
            exp_labels:
              severity: "warning"
              runbook_url: "https://github.com/operator-framework/operator-sdk/tree/master/testdata/go/v4/monitoring/memcached-operator/docs/monitoring/runbooks/MemcachedDeploymentSizeUndesired.md"
      - eval_time: 5m
        alertname: MemcachedOperatorDown
        exp_alerts:
          - exp_annotations:
              description: "No running memcached-operator pods were detected in the last 5 min."
            exp_labels:
              severity: "critical"
              runbook_url: "https://github.com/operator-framework/operator-sdk/tree/master/testdata/go/v4/monitoring/memcached-operator/docs/monitoring/runbooks/MemcachedOperatorDown.md"
      # it must not trigger before 15m
      - eval_time: 14m
        alertname: MemcachedDeploymentSizeUndesired
        exp_alerts: [ ]
      - eval_time: 14m
        alertname: MemcachedOperatorDown
        exp_alerts: [ ]
      # it must trigger after 15m
      - eval_time: 15m
        alertname: MemcachedDeploymentSizeUndesired
        exp_alerts:
          - exp_annotations:
              description: "Memcached-sample deployment size was not as desired more than 3 times in the last 5 minutes."
            exp_labels:
              severity: "warning"
              runbook_url: "https://github.com/operator-framework/operator-sdk/tree/master/testdata/go/v4/monitoring/memcached-operator/docs/monitoring/runbooks/MemcachedDeploymentSizeUndesired.md"
      - eval_time: 15m
        alertname: MemcachedOperatorDown
        exp_alerts:
          - exp_annotations:
              description: "No running memcached-operator pods were detected in the last 5 min."
            exp_labels:
              severity: "critical"
              runbook_url: "https://github.com/operator-framework/operator-sdk/tree/master/testdata/go/v4/monitoring/memcached-operator/docs/monitoring/runbooks/MemcachedOperatorDown.md"
